# Examen <span style="color: #FFA500">Core, funciones incorporadas y strings</span>

**Instrucciones**

* Solamente puedes escribir código entre los comentarios ```### BEGIN SOLUTION``` y ```### END SOLUTION```

## Ejercicio 1 (1 punto)



Crea, inicializa y opera con tres variables ```x, y, z``` de manera que ```z``` valga 30 al terminar la ejecución de la celda.

<!-- grade_id: cell-a54400a1ecd4cad6, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
x = 10
y = 20
z = x + y
### END SOLUTION

print(z)
```

Output:
```
30
```

<!-- grade_id: cell-01b9230322b8f61f, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

assert z == 30

### END HIDDEN TESTS
```

## Ejercicio 2 (1 punto)



* Crea una variable ```x``` e inicialízala con el valor ```float``` 10.5


<!-- grade_id: cell-8fb40d6b5125dab5, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

x = 10.5

### END SOLUTION

print(x)
```

Output:
```
10.5
```

<!-- grade_id: cell-8808493c24993de7, points: 0.5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

assert x == 10.5
assert type(x) == float

### END HIDDEN TESTS
```

* Conviértela, asignándola a sí misma, a ```int``` utilizando la función al efecto de Python.


<!-- grade_id: cell-5522c2752130898c, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

x = int(x)

### END SOLUTION

print(x)
```

Output:
```
10
```

<!-- grade_id: cell-5995b5bc9367119a, points: 0.5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

assert x == 10
assert type(x) == int

### END HIDDEN TESTS
```

## Ejercicio 3 (1 punto)



* Sustitye la vocal ```a``` del texto almacenado en ```x``` por un guión bajo ( ```_``` )

<!-- grade_id: cell-07c99e02b3091a49, points: 0, grade: false, solution: true -->
```python
# Ref.: https://es.wikipedia.org/wiki/Gu%C3%ADa_del_autoestopista_gal%C3%A1ctico_(novela)
x = """
La Guía del autoestopista galáctico (The Hitchhiker's Guide to the Galaxy en inglés) es una novela de ciencia ficción de Douglas Adams publicada en octubre de 1979. Es el primer libro de seis que conforman la famosa serie de la Guía del autoestopista galáctico, que fue originalmente una radiocomedia escrita por el mismo autor en 1978.

El título del libro fue inspirado en la guía Hitch-hiker's Guide to Europe para viajar en autoestop por Europa. La saga de la Guía del autoestopista galáctico ha sido difundida de diferentes maneras, incluyendo una serie de televisión, un videojuego, un cómic y hasta una dramatización teatral, convirtiendo a su autor en una figura influyente en el campo de la ciencia ficción y la comedia.

El libro vendió 250 000 copias en los primeros tres meses tras ser publicado. Hasta 2018 había vendido 16 millones de ejemplares.
"""

### BEGIN SOLUTION

x = x.replace('a', '_')

### END SOLUTION

print(x)
```

Output:
```
L_ Guí_ del _utoestopist_ g_láctico (The Hitchhiker's Guide to the G_l_xy en inglés) es un_ novel_ de cienci_ ficción de Dougl_s Ad_ms public_d_ en octubre de 1979. Es el primer libro de seis que conform_n l_ f_mos_ serie de l_ Guí_ del _utoestopist_ g_láctico, que fue origin_lmente un_ r_diocomedi_ escrit_ por el mismo _utor en 1978.

El título del libro fue inspir_do en l_ guí_ Hitch-hiker's Guide to Europe p_r_ vi_j_r en _utoestop por Europ_. L_ s_g_ de l_ Guí_ del _utoestopist_ g_láctico h_ sido difundid_ de diferentes m_ner_s, incluyendo un_ serie de televisión, un videojuego, un cómic y h_st_ un_ dr_m_tiz_ción te_tr_l, convirtiendo _ su _utor en un_ figur_ influyente en el c_mpo de l_ cienci_ ficción y l_ comedi_.

El libro vendió 250 000 copi_s en los primeros tres meses tr_s ser public_do. H_st_ 2018 h_bí_ vendido 16 millones de ejempl_res.
```

<!-- grade_id: cell-67f6b0b3b0a6dd9e, points: 0.5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

assert x.count('a') == 0

assert x == """
L_ Guí_ del _utoestopist_ g_láctico (The Hitchhiker's Guide to the G_l_xy en inglés) es un_ novel_ de cienci_ ficción de Dougl_s Ad_ms public_d_ en octubre de 1979. Es el primer libro de seis que conform_n l_ f_mos_ serie de l_ Guí_ del _utoestopist_ g_láctico, que fue origin_lmente un_ r_diocomedi_ escrit_ por el mismo _utor en 1978.

El título del libro fue inspir_do en l_ guí_ Hitch-hiker's Guide to Europe p_r_ vi_j_r en _utoestop por Europ_. L_ s_g_ de l_ Guí_ del _utoestopist_ g_láctico h_ sido difundid_ de diferentes m_ner_s, incluyendo un_ serie de televisión, un videojuego, un cómic y h_st_ un_ dr_m_tiz_ción te_tr_l, convirtiendo _ su _utor en un_ figur_ influyente en el c_mpo de l_ cienci_ ficción y l_ comedi_.

El libro vendió 250 000 copi_s en los primeros tres meses tr_s ser public_do. H_st_ 2018 h_bí_ vendido 16 millones de ejempl_res.
"""

### END HIDDEN TESTS
```

* Sustitye las dos primeras vocales ```e``` del texto almacenado en ```x``` resultado de la celda anterior por un guión bajo ( ```_``` )

<!-- grade_id: cell-94b29237334b1bec, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

x = x.replace('e', '_', 2)

### END SOLUTION

print(x)
```

Output:
```
L_ Guí_ d_l _uto_stopist_ g_láctico (The Hitchhiker's Guide to the G_l_xy en inglés) es un_ novel_ de cienci_ ficción de Dougl_s Ad_ms public_d_ en octubre de 1979. Es el primer libro de seis que conform_n l_ f_mos_ serie de l_ Guí_ del _utoestopist_ g_láctico, que fue origin_lmente un_ r_diocomedi_ escrit_ por el mismo _utor en 1978.

El título del libro fue inspir_do en l_ guí_ Hitch-hiker's Guide to Europe p_r_ vi_j_r en _utoestop por Europ_. L_ s_g_ de l_ Guí_ del _utoestopist_ g_láctico h_ sido difundid_ de diferentes m_ner_s, incluyendo un_ serie de televisión, un videojuego, un cómic y h_st_ un_ dr_m_tiz_ción te_tr_l, convirtiendo _ su _utor en un_ figur_ influyente en el c_mpo de l_ cienci_ ficción y l_ comedi_.

El libro vendió 250 000 copi_s en los primeros tres meses tr_s ser public_do. H_st_ 2018 h_bí_ vendido 16 millones de ejempl_res.
```

<!-- grade_id: cell-3eba07b18619d384, points: 0.5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

assert x == """
L_ Guí_ d_l _uto_stopist_ g_láctico (The Hitchhiker's Guide to the G_l_xy en inglés) es un_ novel_ de cienci_ ficción de Dougl_s Ad_ms public_d_ en octubre de 1979. Es el primer libro de seis que conform_n l_ f_mos_ serie de l_ Guí_ del _utoestopist_ g_láctico, que fue origin_lmente un_ r_diocomedi_ escrit_ por el mismo _utor en 1978.

El título del libro fue inspir_do en l_ guí_ Hitch-hiker's Guide to Europe p_r_ vi_j_r en _utoestop por Europ_. L_ s_g_ de l_ Guí_ del _utoestopist_ g_láctico h_ sido difundid_ de diferentes m_ner_s, incluyendo un_ serie de televisión, un videojuego, un cómic y h_st_ un_ dr_m_tiz_ción te_tr_l, convirtiendo _ su _utor en un_ figur_ influyente en el c_mpo de l_ cienci_ ficción y l_ comedi_.

El libro vendió 250 000 copi_s en los primeros tres meses tr_s ser public_do. H_st_ 2018 h_bí_ vendido 16 millones de ejempl_res.
"""

### END HIDDEN TESTS
```

## Ejercicio 4 (1 punto)



Calcula el área de una elipse con los siguientes valores para su radio mayor (```r_a```) y radio menor (```r_b```):
* ```r_a``` vale 10
* ```r_b``` vale 5

$A = \pi*r_a*r_b$

<!-- grade_id: cell-fe8faf9c96779430, points: 0, grade: false, solution: true -->
```python
import math

r_a = 10 # Asigna un valor a r_a
r_b = 5 # Asigna un valor a r_b

### BEGIN SOLUTION

A = math.pi * r_a * r_b

### END SOLUTION

print(A)
```

Output:
```
157.07963267948966
```

<!-- grade_id: cell-7d9f6782826372ca, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

assert A == 157.07963267948966

### END HIDDEN TESTS
```

## Ejercicio 5 (1 punto)



Calcula el rendimiento diario de Apple entre los días 28/09/2023 y 29/09/2023 y almacénalo en la variable ```rentab```
* VL del 28/09/2023 es 169.34
* VL del 29/09/2023 es 172.02

$rentab = \frac{(vl_2 - vl_1)}{vl_1}$

<!-- grade_id: cell-ca005f751a58e53c, points: 0, grade: false, solution: true -->
```python
vl_1 = 169.34 # Asigna un valor a vl_1
vl_2 = 172.02 # Asigna un valor a vl_2

### BEGIN SOLUTION

rentab = (vl_2 - vl_1) / vl_1

### END SOLUTION

print(rentab)
```

Output:
```
0.015826148576827723
```

<!-- grade_id: cell-04d3b6b8c638587e, points: 0.5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

assert rentab == 0.015826148576827723

### END HIDDEN TESTS
```

Muestra el valor de la rentabilidad calculada en una cadena de texto.

La cadena de texto final tiene que ser la siguiente: ```Apple obtuvo un rendimiento diario entre los días 28/09/2023 y 29/09/2023 de un 1.58%``` y estar almacenada en la variable ```mensaje```. Debes usar las variables ```fecha_vl_1```, ```fecha_vl_2``` y ```rentab``` para construir la cadena de texto final.

<!-- grade_id: cell-6e5cf3170070bfd5, points: 0, grade: false, solution: true -->
```python
from datetime import datetime

fecha_vl_1 = datetime(2023, 9, 28) # Asigna un valor a fecha_vl_1 utilizando datetime
fecha_vl_2 = datetime(2023, 9, 29) # Asigna un valor a fecha_vl_2 utilizando datetime

### BEGIN SOLUTION

mensaje = f"Apple obtuvo un rendimiento diario entre los días {fecha_vl_1:%d/%m/%Y} y {fecha_vl_2:%d/%m/%Y} de un {rentab*100:.2f}%"

### END SOLUTION

print(mensaje)

```

Output:
```
Apple obtuvo un rendimiento diario entre los días 28/09/2023 y 29/09/2023 de un 1.58%
```

<!-- grade_id: cell-9834a4d32990072c, points: 0.5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

assert mensaje == "Apple obtuvo un rendimiento diario entre los días 28/09/2023 y 29/09/2023 de un 1.58%"

### END HIDDEN TESTS
```

## Ejercicio 6 (1 punto)



Supongamos que tienes una cuenta de ahorros con un interés anual del 2% y quieres saber cuánto dinero tendrás después de un año si depositas 1000 euros.

$VF = VP \times (1+r)$


* Primero, asigna el valor 1000 a una variable llamada ```saldo_inicial```.

<!-- grade_id: cell-79176afa6cc04914, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

saldo_inicial = 1000

### END SOLUTION

print(saldo_inicial)
```

Output:
```
1000
```

<!-- grade_id: cell-0d1fc5c5ddde7289, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

assert saldo_inicial == 1000

### END HIDDEN TESTS
```

* Segundo, calcula el interés que se acumulará en un año multiplicando el saldo inicial por el porcentaje de interés y dividiéndolo por 100. Almacénalo en una variable llamada ```interes```.

<!-- grade_id: cell-854b9a2ef96b3712, points: 0, grade: false, solution: true -->
```python
# Tasa de interés anual
r = 2

### BEGIN SOLUTION

interes = saldo_inicial * r / 100

### END SOLUTION

print(interes)
```

Output:
```
20.0
```

<!-- grade_id: cell-5c0dc054cced4713, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

assert interes == 20

### END HIDDEN TESTS
```

* Tercero, suma el saldo inicial y el interés para obtener el saldo final después de un año. Almacénalo en la variable ```saldo_final```

<!-- grade_id: cell-a9d612a0cf6e4038, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

saldo_final = saldo_inicial + interes

### END SOLUTION

print(saldo_final)
```

Output:
```
1020.0
```

<!-- grade_id: cell-c1f214d81aac268b, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

assert saldo_final == 1020

### END HIDDEN TESTS
```

* Cuarto, muestra el saldo final con dos decimales y el símbolo del euro. Crea una variable llamada ```mensaje``` que contenga el texto ```El saldo final después de un año es 1020.00€```, usando la variable ```saldo_final``` para mostrar la cifra.

<!-- grade_id: cell-b3b76f1b6c7c3575, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

mensaje = f"El saldo final después de un año es {saldo_final:.2f}€"

### END SOLUTION

print(mensaje)
```

Output:
```
El saldo final después de un año es 1020.00€
```

<!-- grade_id: cell-c539f89938bf6348, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

assert mensaje == "El saldo final después de un año es 1020.00€"

### END HIDDEN TESTS
```

## Ejercicio 7 (1 punto)



Supongamos que quieres comprar un ordenador portátil que cuesta 800 euros, pero solo tienes 500 euros ahorrados. Decides pedir un préstamo al banco por la diferencia con un interés anual del 10% y devolverlo en 18 meses.

* Primero, calcula el importe total a devolver al banco. Almacena este valor en la variable ```total_devolver```

<!-- grade_id: cell-232ce417a49fa0c1, points: 0, grade: false, solution: true -->
```python
r = 0.10 # Tasa de interés anual del 10%
importe_total = 800
importe_ahorros = 500
cuotas = 18

### BEGIN SOLUTION

total_devolver = (importe_total - importe_ahorros) * (1 + (r/12))**cuotas

### END SOLUTION

print(total_devolver)
```

Output:
```
348.3336989406774
```

<!-- grade_id: cell-3190b63d82be068b, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

assert total_devolver == 348.3336989406774

### END HIDDEN TESTS
```

* Segundo, calcula la cuota mensual que tendrás que pagar al banco. Obtén la división entera, es decir, sin decimales, y guárdala en la variable ```cuota``` de tipo entero.

<!-- grade_id: cell-78cabf2db7490ce3, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

cuota = int(total_devolver // cuotas)

### END SOLUTION

print(cuota)
```

Output:
```
19
```

<!-- grade_id: cell-7c6f3b4ef0e04a55, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

assert cuota == 19
assert type(cuota) == int

### END HIDDEN TESTS
```

* Tercero, calcula el resto que te quedará por pagar después de pagar la última cuota. Obtén el módulo, es decir, el residuo de la división entera, y guárdalo en la variable ```resto```.

<!-- grade_id: cell-84cd871694f79e23, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

resto = total_devolver % cuotas

### END SOLUTION

print(resto)
```

Output:
```
6.33369894067738
```

<!-- grade_id: cell-d3efa612f2f29807, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

assert resto == 6.33369894067738
assert type(resto) == float

### END HIDDEN TESTS
```

* Calcula el importe de la última cuota y almacénalo en la variable ```ultima_cuota```
* Muestra un mensaje en el que se indique ```Las primeras 17 cuotas serán de 19€, la última cuota será de 25.33€``` utilizando las variables calculadas anteriormente.

<!-- grade_id: cell-771d1db654ea8489, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

ultima_cuota = cuota + resto
mensaje = f"Las primeras 11 cuotas serán de {cuota}€, la última cuota será de {ultima_cuota:.2f}€"

### END SOLUTION

print(mensaje)
```

Output:
```
Las primeras 11 cuotas serán de 19€, la última cuota será de 25.33€
```

<!-- grade_id: cell-2bf98eb987d8f66d, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

assert ultima_cuota == 25.33369894067738
assert type(ultima_cuota) == float
assert mensaje == "Las primeras 11 cuotas serán de 19€, la última cuota será de 25.33€"

### END HIDDEN TESTS
```

## Ejercicio 8 (1 punto)



Supongamos que queremos escribir un programa que pida al usuario que introduzca una cantidad de dinero en euros y luego le muestre el equivalente en dólares.
* Usaremos un tipo de cambio fijo de 1 euro = 1.17 dólares.
* El usuario solo puede introducir números válidos, si no, mostrará un mensaje de error.
* Debes usar las funciones ```round```, ```str.split``` y ```str.isnumeric``` para este ejercicio.

Un ejemplo de salida sería:

Introduzca una cantidad de dinero en euros: 50\
El equivalente en dólares es 58.5

Otro ejemplo de salida sería:

Introduzca una cantidad de dinero en euros: 12.75\
El equivalente en dólares es 14.92

Otro ejemplo de salida sería:

Introduzca una cantidad de dinero en euros: hola\
Error: por favor, introduzca solo números.


<!-- grade_id: cell-f12cf5e5cf8cc824, points: 0, grade: false, solution: true -->
```python
# Definimos la tasa de cambio fija
tasa = 1.17

### BEGIN SOLUTION

def convertir_euros_a_dolares(euros:float, tasa:float)->float:
    """Convierte una cantidad de euros a dólares.

    Args:
        euros (float): cantidad de euros a convertir
        tasa (float): tasa de cambio de euros a dólares

    Returns:
        float: cantidad de dólares equivalente a la cantidad de euros
    """
    # Calcular el equivalente en dólares usando la tasa de cambio
    dolares = euros * tasa

    # Redondear el resultado a dos decimales usando la función round
    dolares = round(dolares, 2)

    return dolares

def comprobar_si_es_numero(numero:str)->bool:
    """Comprueba si una cadena es un número válido.

    Args:
        numero (str): cadena con el número a comprobar

    Returns:
        bool: True si la cadena es un número válido, False en caso contrario
    """
    # Comprobar si el usuario ha introducido solo números
    # Para ello, se puede usar la función str.split para separar la parte entera y la decimal por el punto
    # y luego usar la función str.isnumeric para ver si ambas partes son numéricas
    # Una cadena es numérica si todos los caracteres son numéricos y hay al menos un caracter.
    partes = numero.split(".")
    if len(partes) == 1 and partes[0].isnumeric():
        # El usuario ha introducido un número entero
        valido = True
    elif len(partes) == 2 and partes[0].isnumeric() and partes[1].isnumeric():
        # El usuario ha introducido un número decimal
        valido = True
    else:
        # El usuario ha introducido algo que no es un número
        valido = False
    
    return valido

# Pedimos al usuario que introduzca una cantidad de dinero en euros
euros = input("Introduzca una cantidad de dinero en euros: ")

# Comprobar si el usuario ha introducido solo números
valido = comprobar_si_es_numero(euros)

# Si el usuario ha introducido un número válido, calcular el equivalente en dólares y mostrarlo por pantalla
if valido:
    # Convertir el string euros a un número real
    euros = float(euros)

    # Calcular el equivalente en dólares usando la tasa de cambio
    dolares = convertir_euros_a_dolares(euros, tasa)
    
    # Mostrar el resultado por pantalla
    print("El equivalente en dólares es", dolares)
else:
    # Mostrar un mensaje de error por pantalla
    print("Error: por favor, introduzca solo números.")

### END SOLUTION
```

Output:
```
El equivalente en dólares es 11.7
```

<!-- grade_id: cell-8d2aff0fdcf50720, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

assert convertir_euros_a_dolares(100, 1.17) == 117
assert comprobar_si_es_numero("100") == True
assert comprobar_si_es_numero("hola") == False

### END HIDDEN TESTS
```

## Ejercicio 9 (1 punto)



Crea una lista que se llame ```numeros``` con los números del 1 al 20 en orden descendente.

Ref.:
* https://docs.python.org/3/tutorial/introduction.html#lists
* https://docs.python.org/3/tutorial/datastructures.html
* https://docs.python.org/3/tutorial/controlflow.html#for-statements
* https://docs.python.org/3/tutorial/controlflow.html#the-range-function

<!-- grade_id: cell-6312cf4cac45b78d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

# Opción 1: usando un bucle for
numeros = []
for i in range(20, 0, -1):
    numeros.append(i)

# Opción 2: usando range y una conversión a lista
numeros = list(range(20, 0, -1))

### END SOLUTION
```

<!-- grade_id: cell-c29d92e358eb72bb, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

assert numeros == [20, 19, 18, 17, 16, 15, 14, 13, 12, 11, 10, 9, 8, 7, 6, 5, 4, 3, 2, 1]

### END HIDDEN TESTS
```

## Ejercicio 10 (1 punto)

Crea una función que calcule los 20 primeros números de la sucesión de Fibonacci y los almacene en una variable llamada ```numeros```.

* Los dos primeros números de la sucesión deben ser ```[0, 1]```
* La sucesión que se espera como resultado es ```[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181, 6765, 10946]```

Ref.:
* https://es.wikipedia.org/wiki/Sucesión_de_Fibonacci

<!-- grade_id: cell-20b38745a8afe02b, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def sucesion_fibonacci():
    sucesion = [0, 1]

    for i in range(20):
        sucesion.append(sucesion[-1] + sucesion[-2])
    
    return sucesion

numeros = sucesion_fibonacci()

### END SOLUTION
```

<!-- grade_id: cell-69e12d30e4ae3559, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

assert numeros == [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181, 6765, 10946]

### END HIDDEN TESTS
```